# Week 5 Exercises

_McKinney 3.2_

**Unlike in previous weeks, in the exercises below, you will need to create a function definition from scratch.  I'll provide specific instructions and examples for how it will be used, but you will have to do the work of definiging the whole function.**

---
---

### 25.1 Longest String

Write a function called `longest (L)` that takes as its only parameter a list of strings.  Your function needs to find the longest of those strings and return the position number of that longest string.  An example is provided below:

The longest word in that list of strings is "birthday" in position #1, so `longest(strings)` should return 1.
```
>>> strings = ['happy', 'birthday', 'to', 'me']
>>> longest(strings)
1
```

Be sure to include a docstring.  Include test cases in your docstring if you want to.

```python
# in class example
def longest(my_list):
    '''(list) -> int
    Find longest item in a list of strings and return the index for that string. 
    If there are multiple strings with the same length, return the index of the first of those longest strings.
    
    >>> longest(['happy', 'birthday', 'to', 'me'])
    1
    
    >>> longest(['abc', 'def', 'ghi'])
    0
    
    '''
    # initialize intermediate variables 
    # length at -1, as any value should be greater than that if there is a value present
    max_length = -1
    
    # index of the longest value - initalize at None
    max_index = None
    
    # demonstrate handling syntax
    for index, item in enumerate(my_list):
        try:
            if len(item) > max_length:
                max_length = len(item)
                max_index = index
        # If length of the item cannot be calculated, pass and move to next item - this avoids TypeError
        except:
            pass
            
    # at the end of the for loop, return the resulting index of the max length
    return max_index
```

<!-- grade_id: longest-answer, points: 0, grade: false, solution: true -->
```python
### BEGIN SOLUTION

def longest(strings):
    """(list) -> int
    This function finds the longest string in the list of input strings
    and returns the position number of that longest string.
    
    >>> longest(['happy', 'birthday', 'to', 'me'])
    1
    
    >>> longest(['one', 'two', 'three', 'four'])
    2
    """
    
    longest_pos = 0
    longest_len = 0
    
    print("Position  Item         Longest So Far")
    print("--------- ------------ --------------")
    for position, item in enumerate(strings):
        
        if len(item) > longest_len:
            longest_pos = position
            longest_len = len(item)
        print("{0:9} {1:12} {2} : {3}".format(position, item, longest_pos, longest_len))
            
    return longest_pos

### END SOLUTION
```

```python
longest(['happy', 'birthday', 'to', 'me'])
```

Output:
```
1
```

<!-- grade_id: longest-tests, points: 10, grade: true, solution: false -->
```python
assert longest("happy birthday to me".split(" ")) == 1
assert longest("enjoy class".split(" ")) == 0
assert longest(['when','what','where','how','who']) == 2
```

Output:
```
Position  Item         Longest So Far
--------- ------------ --------------
        0 happy        0 : 5
        1 birthday     1 : 8
        2 to           1 : 8
        3 me           1 : 8
Position  Item         Longest So Far
--------- ------------ --------------
        0 enjoy        0 : 5
        1 class        0 : 5
Position  Item         Longest So Far
--------- ------------ --------------
        0 when         0 : 4
        1 what         0 : 4
        2 where        2 : 5
        3 how          2 : 5
        4 who          2 : 5
```

```python
import doctest
doctest.run_docstring_examples(longest, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    longest(['happy', 'birthday', 'to', 'me'])
Expecting:
    1
ok
Trying:
    longest(['abc', 'def', 'ghi'])
Expecting:
    0
ok
```

### 25.2 Celsius to Farenheit

Write a function called f_to_c() that converts a given temperature in degrees Farenheit to degrees Celsius.  If you don't recall that conversion, it is:

$ temp_c = \frac{5}{9} \times (temp_f - 32) $

Make sure that you code is well documented using the DocString examples, and that your code includes tests for 212f, 32f, and 98.6f.

```python
### BEGIN SOLUTION - consider all exception handling possible (i.e. non-numeric input, etc.)
# Define the function f_to_c() which takes one numeric input
def f_to_c(temp_f):
    '''(int or float) -> float
    This function takes a temperature in degrees Farenheit (temp_f) and converts it to a temperature in degrees Celcius (temp_c) 
    using the following formula:
    temp_c = (5/9) * (temp_f - 32)
    
    Any input other than an integer or a float will raise a TypeError: "Please enter a number (float or integer)"
    
    >>> f_to_c(212)
    100.0
    
    >>> f_to_c(32)
    0.0
    
    >>> f_to_c(98.6)
    37.0
    '''
    
    # Check if the input is not a number (integer or float)
    if (type(temp_f) != int) or (type(temp_f) != float):
        raise TypeError("Please enter a number (float or integer)") 
    
    # If TypeError is not raised, input must be either a float or an integer
    # Proceed with calculation - convert temp_f to temp_c per formula above
    temp_c = (5/9) * (temp_f - 32)
    
    # return temp_c
    return temp_c

### END SOLUTION
```

```python
### BEGIN SOLUTION
def f_to_c(F):
    """(float) -> float
    
    >>> f_to_c(212)
    100.0
        
    >>> f_to_c(32)
    0.0

    >>> f_to_c(98.6)
    37.0
    """
    return (5/9) * (F-32)

### END SOLUTION
```

```python
assert f_to_c(212) == 100.0
assert f_to_c(32) == 0.0
assert f_to_c(98.6) == 37.0
```

```python
f_to_c(212)
```

Output:
```
100.0
```

```python
import doctest
doctest.run_docstring_examples(f_to_c, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    f_to_c(212)
Expecting:
    100.0
ok
Trying:
    f_to_c(32)
Expecting:
    0.0
ok
Trying:
    f_to_c(98.6)
Expecting:
    37.0
ok
```

### 25.3 Computing Length of Stay

For this problem, we have a collection of patient enounter data stored as a Python dictionary.  The `key` for the dictionary is the **encounter ID**, a code that starts with the letter `E` followed by four numbers.  The value associated with each encounter ID is another Python dictionary.  This "inner" dictionary holds three items: admit date, primary diagnosis, and discharge date.  See the example in the code below.

You need to write a length of stay function that computes the length of stay, in whole days, between the admit date and discharge date.  However, if the diagnosis is "Observation" then the length of stay should always be returned as 0 regardless of the admit and discharge dates.

You will find it handy to refer to this example here on how to calculate the number of days between two dates: https://stackoverflow.com/questions/151199/how-to-calculate-number-of-days-between-two-given-dates

```python
from datetime import date

### BEGIN SOLUTION
# Define the function los that accepts 3 parameters: admit (a date), diagnosis (string), and discharge (a date)
def los(admit, discharge, diagnosis):
    # Doc String
    ''' (date, date, string) -> int
    Define the function los that accepts 3 parameters: 
    admit (a date), discharge (a date), and diagnosis (string)
    
    The function returns an integer representing the length of stay in whole days between admit 
    and discharge. If the diagnosis is "Observation" then los should return 0 regardless of dates.
    
    TypeError is raised if any of the input values are not appropriate for the parameter
        
    >>> los(date(2019,1,3), date(2019,1,8), "COPD")
    5
    
    >>> los(date(2020,1,5), date(2020,1,8), "Observation")
    0
    '''
    # Check the parameters for correct type. Raise TypeError for any incongruent parameters
    # This does not function as expected, not sure why
    
    # If a date is not submitted, this error is not raised. However, the computation cannot occur
    # and a different error is raised.
    if type(admit) != datetime.date:
        raise TypeError("Please submit an admission date as a date() object. See help for details.")
    
    # If a date is not submitted, this error is not raised. However, the computation cannot occur
    # and a different error is raised.
    if type(discharge) != datetime.date:
        raise TypeError("Please submit a discharge date as a date() object. See help for details.")
    
    # This should work but doesn't and no alternate error is raised
    if type(diagnosis) != str:
        raise TypeError("Please submit a diagnosis as a string. See help for details.")
    
    # Initialize output variable los for the date difference
    los = -1
    # Check the diagnosis type for "Observation"
    if diagnosis == "Observation":
        # If diagnosis is "Observation", then los = 0
        los = 0
    else:
        # If the diagnosis is not "Observation", calculate difference between dates in days
        delta = discharge - admit
        # Set return variable to the number of days
        los = delta.days
    
    # Return the length of stay
    return los
### END SOLUTION
```

```python
def los(admit, discharge, diagnosis):
    """(date, str, date)->int
    Computes length of stay in days.  If diagnosis is "Observation" then set the length of stay to 0.

    >>> los(date(2021,1,3), date(2021,1,8), 'COPD')
    5
    
    >>> los(date(2021,1,5), date(2021,1,9), 'Hypertension')
    4
    
    >>> los(date(2021,1,23), date(2021,1,25), 'Observation')
    0
    """
    if diagnosis == "Observation":
        return 0
    else:
        return (discharge - admit).days
```

```python
import doctest
doctest.run_docstring_examples(los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    los(date(2021,1,3), date(2021,1,8), 'COPD')
Expecting:
    5
ok
Trying:
    los(date(2021,1,5), date(2021,1,9), 'Hypertension')
Expecting:
    4
ok
Trying:
    los(date(2021,1,23), date(2021,1,25), 'Observation')
Expecting:
    0
ok
```

```python

```

```python
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}


for e, v in encounters.items():
    print(los(v['admit'],v['discharge'],v['diagnosis']))
```

Output:
```
5
4
1
0
0
```

---
---

### 25.4 Average Length of Stay

Create a function called `average_los` that returns the average of the LOS for the encounters in the provided dictionary. The encounters must be a dictionary that contains dictionaries that each contain at least an admit date and a discharge date.
    
If the length of stay of any individual encounter is 0, then it will not be counted toward the average.

```python
test1 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

test2 = {
        "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) }
}

### BEGIN SOLUTION
# Define the function average_los
def average_los(encounter_dict):
    """ (dict) -> float
    Given a dictionary of encounters containing an admission date and discharge date, 
    calculate the length of stay for each encounter and then calculate the average length of stay
    across all encounters in the dictionary. 
    If the length of stay for any encounter is 0, it will not be counted toward the average
    
    >>> average_los(test1)
    5.0
    """
    # Check for the appropriate parameter input
    if type(encounter_dict) != dict:
        raise TypeError("Please submit a dictionary. See help for details.")
    
    # Initialize a list for calculated length of stay results
    los_list = []
    
    # Initialize a variable to return
    avg_los = -1
    
    # For each key, value pair in the dictionary
    for key, value in encounter_dict.items():
        # Calculate the length of stay
        los = (value['discharge'] - value['admit']).days
        
        # if the length of stay is greater than 0, append the length of stay to a list
        if los > 0:
            los_list.append(los)
        # if length of stay is 0, pass
        if los == 0:
            pass
        # If length of stay is a negative number, raise value error
        if los <0:
            raise ValueError("The calculated length of stay is a negative number. Check the values of encounter {}".format(key))
    
    # Given the list of calculated length of stay values, calculate the average
    # Sum the values of the list and divide by the length of the list
    avg_los = sum(los_list)/len(los_list)
    
    # Return the average length of stay
    return avg_los
### END SOLUTION
```

```python
# testing function average_los
testing = {
        "E1234": { "admit": date(2019,1,8), "diagnosis": "COPD", "discharge": date(2019,1,3) }
}

# tested typeError - worked as expected
#average_los(1)

# testing ValueError - should produce negative los and thus error - worked as expected
# average_los(testing)
```

```python
from datetime import date
encounters = { 
    "E1234": { "admit": date(2019,1,3), "diagnosis": "COPD", "discharge": date(2019,1,8) },
    "E8342": { "admit": date(2019,1,5), "diagnosis": "Hypertension", "discharge": date(2019,1,9) },
    "E9231": { "admit": date(2019,1,12), "diagnosis": "Anxiety", "discharge": date(2019,1,13) },
    "E8333": { "admit": date(2019,1,15), "diagnosis": "Observation", "discharge": date(2019,1,16) },
    "E3342": { "admit": date(2019,1,4), "diagnosis": "Anxiety", "discharge": date(2019,1,4)}
}

assert(average_los(encounters)) == 2.75
```

```python
import doctest
doctest.run_docstring_examples(average_los, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    average_los(test1)
Expecting:
    5.0
ok
```

---

### 25.5 Celsius to Farenheit

Write your own function called `c_to_f` that converts degrees Celsius to degrees Farenheit.  Include in your solution a series of doc tests that can verify the conversion using inputs of 100, 0, and 37 degrees Celsius

```python
### BEGIN SOLUTION
# Create a function c_to_f that converts a temperature in degrees Celsius 
# to degrees Farenheit
def c_to_f(temp_c):
    """(float or int) -> float
    This function converts a temperature in degrees Celsius to degrees Farenheit
    Using the following equation:
    temp_f = (temp_c *(9/5)) + 32
    
    >>> c_to_f(100.0)
    212.0
    
    >>> c_to_f(0.0)
    32.0
    
    >>> c_to_f(37.0)
    98.6
    """
    
    # Test if the input is a float or integer
    if (type(temp_c) != int) or (type(temp_c) != float):
        raise TypeError("Please enter a number (float or integer)") 
    
    # If TypeError is not raised, the input temperature must be a number. Proceed.
    # Convert the temperature using the equation above
    temp_f = (temp_c *(9/5)) + 32
    
    # return the temperature in farenheit
    return temp_f

### END SOLUTION
```

```python
type(100)
```

Output:
```
int
```

```python
import doctest
doctest.run_docstring_examples(c_to_f, globals(), verbose=True)
```

Output:
```
Finding tests in NoName
Trying:
    c_to_f(100.0)
Expecting:
    212.0
**********************************************************************
File "__main__", line 10, in NoName
Failed example:
    c_to_f(100.0)
Exception raised:
    Traceback (most recent call last):
      File "/opt/tljh/user/lib/python3.6/doctest.py", line 1330, in __run
        compileflags, 1), test.globs)
      File "<doctest NoName[0]>", line 1, in <module>
        c_to_f(100.0)
      File "<ipython-input-71-587005f1af12>", line 22, in c_to_f
        raise TypeError("Please enter a number (float or integer)")
    TypeError: Please enter a number (float or integer)
Trying:
    c_to_f(0.0)
Expecting:
    32.0
**********************************************************************
File "__main__", line 13, in NoName
Failed example:
    c_to_f(0.0)
Exception raised:
    Traceback (most recent call last):
      File "/opt/tljh/user/lib/python3.6/doctest.py", line 1330, in __run
        compileflags, 1), test.globs)
  
```

---

## Check your work above

If you didn't get them all correct, take a few minutes to think through those that aren't correct.


## Submitting Your Work

In order to submit your work, you'll need to use the `git` command line program to **add** your homework file (this file) to your local repository, **commit** your changes to your local repository, and then **push** those changes up to github.com.  From there, I'll be able to **pull** the changes down and do my grading.  I'll provide some feedback, **commit** and **push** my comments back to you.  Next week, I'll show you how to **pull** down my comments.

To run through everything one last time and submit your work:
1. Use the `Kernel` -> `Restart Kernel and Run All Cells` menu option to run everything from top to bottom and stop here.
2. Follow the instruction on the prompt below to either ssave and submit your work, or continue working.

If anything fails along the way with this submission part of the process, let me know.  I'll help you troubleshoort.

---

```python
a=input('''
Are you ready to submit your work?
1. Click the Save icon (or do Ctrl-S / Cmd-S)
2. Type "yes" or "no" below
3. Press Enter

''')

if a=='yes':
    !git add week05_assignment_2.ipynb
    !git commit -a -m "Submitting the week 5 programming exercises"
    !git push
else:
    print('''
    
OK. We can wait.
''')
```


---

If the message above says something like _Submitting the week 3 review exercises_ or _Everything is up to date_, then your work was submitted correctly.

```python

```